```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

### Importing Libraries and Magic Commands

In CSCI 3022, we will be using common Python libraries to help us process data. By convention, we import all libraries at the very top of the notebook. There are also a set of standard aliases that are used to shorten the library names. Below are some of the libraries that you may encounter throughout this assignment, along with their respective aliases.

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import plotly.express as px
from csci3022_utils import *

```

<a id='verytop'></a>

# Homework 2: Prerequisite Review Part 2 & Exploratory Data Analysis

## Due Date: ~~Thursday, Jan 25th,~~ Friday, Jan 26th 11:59 PM on Gradescope


### Detailed Submission Instructions Are Provided at the end of this Notebook







## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Piazza!



## Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1a | 3 | autograded
Question 1b | 4 | manual
Question 1c | 3 | manual
Question 2a | 4 | manual
Question 2b | 6 | manual
Question 3a | 3 | auto
Question 3b | 3 | manual
Question 3c | 1 | autograded
Question 3d| 4 | manual
Question 4a | 8 | autograded and manual
Question 4b | 4 | autograded
Question 4c | 4 | autograded
Question 4d | 3 | manual
|Total | 50

## Recommended Readings for this HW:

 - [Prerequisite Review Resources posted on the  Modules in Canvas](https://canvas.colorado.edu/courses/101142/pages/prerequisite-review-calculus-and-discrete-structures?module_item_id=5144193) 
 - [Learning Data Science Chapter 6 - All (Subsetting, Aggregating, Joining, Transforming)](https://learningds.org/ch/06/pandas_intro.html)
 - [Learning Data Science Section 8.6 (Table Shape and Granularity)](https://learningds.org/ch/08/files_granularity.html#table-shape-and-granularity)
 


<a id='top'></a>
---
**Shortcuts:**   [Question 2](#p2) | [Question 3](#p3) | [Question 4](#p4) | 
---

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />




## <span style='color:Red'>  Question 1: Prerequisites:  Discrete Structures Review ##

You will need a solid understanding of the following probability concepts from Discrete Structures to succeed in this course. 

See the [Prerequisite Review Resources](https://canvas.colorado.edu/courses/101142/pages/prerequisite-review-calculus-and-discrete-structures?module_item_id=5144193) posted on Canvas if you need to review any of these key concepts.


### Question 1a) ####  

A coin is flipped 10 times.  How many possible outcomes have exactly 3 heads?  

Use code to enter your answer in the cell below.  

```python
q1a_answer = ...
q1a_answer
```

```python
grader.check("q1a")
```

<!-- BEGIN QUESTION -->

#### Question 1b) ####  

What is the probability that if I roll two 6-sided dice they add up to **at most** $9$? Use LaTeX (not code) in the cell directly below to show all of your steps and fully justify your answer.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

1b Answer Check).  To check your final answer to 1b, enter the answer you came up with (just the number) in the cell below. Note that this is just a built-in public test so you can check your work and determine if you are on the right track.  To receive credit on this problem you must show all steps in part 1b above using LaTeX and fully justifying your answer using correct mathematical notation.  

```python
q1b_answer = ...
```

```python
grader.check("q1bcheck")
```

<!-- BEGIN QUESTION -->

#### Question 1c) #### 

Suppose you uncharacteristically show up to a quiz completely unprepared.  The quiz has 10 questions, each with 5 multiple choice options. You decide to guess each answer in a completely random way.  What is the probability that you get exactly 3 questions correct?  Use Markdown and LaTeX (not code) in the cell directly below to show all of your steps and fully justify your answer. 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

1c Final Answer Check).  To check your final answer to 1c, enter the final answer you came up with in the cell below. Do not round your answer. Note that this is just a built-in public test so you can check your work and determine if you are on the right track.  To receive credit on this problem you must show all steps in part 1c above using LaTeX and fully justifying your answer using correct mathematical notation.  

```python
q1c_answer = ...
```

```python
grader.check("q1ccheck")
```

[Back to top](#top)

<a id='p2'></a>

## <span style='color:Red'>  Question 2: Prerequisites: Calculus Review ##

You will need a solid understanding of the following concepts from Calculus 1 and 2 to succeed in this course. See the [Prerequisite Review Resources](https://canvas.colorado.edu/courses/101142/pages/prerequisite-review-calculus-and-discrete-structures?module_item_id=5144193) posted on Canvas if you need to review any of these key concepts.


### Preliminary: Sums ###

Here's a recap of some basic algebra written in sigma notation. The facts are all just applications of the ordinary associative and distributive properties of addition and multiplication, written compactly and without the possibly ambiguous "...". But if you are ever unsure of whether you're working correctly with a sum, you can always try writing $\sum_{i=1}^n a_i$ as $a_1 + a_2 + \cdots + a_n$ and see if that helps.

- You can use any reasonable notation for the index over which you are summing, just as in Python you can use any reasonable name in `for name in list`. Thus $\sum_{i=1}^n a_i = \sum_{k=1}^n a_k$.
- $\sum_{i=1}^n (a_i + b_i) = \sum_{i=1}^n a_i + \sum_{i=1}^n b_i$
- $\sum_{i=1}^n d = nd$
- $\sum_{i=1}^n (ca_i + d) = c\sum_{i=1}^n a_i + nd$


<!-- BEGIN QUESTION -->

### Question 2a) #### 

We commonly use sigma notation to compactly write the definition of the arithmetic mean (commonly known as the average):

$$\bar{x} = \frac{1}{n}\left(x_1+x_2+ ... + x_n \right) = \frac{1}{n}\sum_{i=1}^n x_i$$



The $i$th *deviation from average* is the difference $x_i - \bar{x}$. Prove that the sum of all these deviations is 0 that is, prove that $\sum_{i=1}^n (x_i - \bar{x}) = 0$ (write your full solution in the box directly below showing all steps and using LaTeX).


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2b) ####  

Let $x_1, x_2, \ldots, x_n$ be a list of numbers. You can think of each index $i$ as the label of a household, and the entry $x_i$ as the annual income of Household $i$. 

Consider the function  $$f(c) = \frac{1}{n} \sum_{i=1}^n (x_i-c)^2$$


In this scenario, suppose that our data points $x_1, x_2, \ldots, x_n$ are fixed and that $c$ is the only variable.

Using calculus, determine the value of $c$ that minimizes $f(c)$.  You must use calculus to justify that this is indeed a minimum, and not a maximum.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

[Back to top](#top)

<a id='p3'></a>

## <span style='color:Red'>  Question 3).  Applying Those Prereqs:  A Maximum Likelihood Estimate ##

In this problem we're going to apply your calculus and discrete math prerequisite knowledge to introduce a data science concept called a maximum likelihood estimate.

Data scientists use coin tossing as a visual image for sampling at random with replacement from a binary population.

### Question 3a) ####
A coin that lands heads with chance 0.7 is tossed six times. What is the chance of the sequence HHHTHT? Assign your answer to the variable `p_HHHTHT`.

```python
p_HHHTHT = ...
p_HHHTHT
```

```python
grader.check("3a")
```

### Question 3b) ####
I have a coin that lands heads with an unknown probability $p$. 

Suppose I toss it 10 times and get the sequence TTTHTHHTTH.

If you toss this coin 10 times, the chance that you get the sequence above is a function of $p$. That function is called the *likelihood* of the sequence TTTHTHHTTH, so we will call it $L(p)$.



<!-- BEGIN QUESTION -->

What is $L(p)$ (i.e. the likelihood) for the sequence TTTHTHHTTH?

Enter your answer below by setting the `likelihood` variable equal to the correct function.

(For example `likelihood = sin(p)+2p`, althought that is definitely an incorrect answer!)

Then run the code below to plot the likelihood function.



```python
#At the top of the notebook we already imported a useful plotting module, matplotlib with alias plt

p = np.linspace(0, 1, 100) 
#This creates an array of 100 p-values equally spaced between 0 and 1

likelihood = ...
#Define the likelihood function above

plt.plot(p, likelihood, lw=2, color='darkblue') 
#This plots the likelihood function  

plt.xlabel('$p$')
#This labels the x axis

plt.ylabel('$L(p)$')
#This labels the y-axis

plt.title('Likelihood of TTTHTHHTTH');
#This titles the plot
```

<!-- END QUESTION -->

### Question 3c) ####

The value $\hat{p}$ at which the likelihood function attains its maximum is called the *maximum likelihood estimate* (MLE) of $p$. Among all values of $p$, it is the one that makes the observed data most likely.

Using only your plot above, what is the value of $\hat{p}$?   


```python
p_hat = ...
```

```python
grader.check("3c")
```

<!-- BEGIN QUESTION -->

### Question 3d) ####

Notice the value you found graphically for $\hat{p}$ above also intuitively makes sense because it is also the observed proportion of heads in the given sequence TTTHTHHTTH.


Let's prove what you observed graphically above.  That is, let's use calculus to find $\hat{p}$.  

But **wait before you start trying to find the value $p$ where $L'(p)=0$ (trust us, the algebra is not pretty...)**

USEFUL TIP: 
The value $\hat{p}$ at which the function $L(p)$ attains its maximum is the same as the value at which the function $\ln(L(p))$ attains its maximum. 

This tip is hugely important in data science because many probabilities are products and the natural log function `ln` function turns products into sums.  It's **much simpler to take derivatives of a sum** than a product.


Thus, to find the value $p$ where $L'(p)=0$: 
 - Take the natural log `ln` of L(p)
 - Use properties of logs to rewrite products in `ln(L(p))` as sums
 - Take the derivative of this rewritten version of `ln(L(p))`
 - Solve `ln(L(p))=0` for p
 - You should get the same answer that you found graphically above.

You don't have to check that the value you've found produces a max and not a min – we'll spare you that step.

Show all steps in the cell below using Markdown and LaTeX





_Type your answer here, replacing this text._

<!-- END QUESTION -->

[Back to top](#top)

<a id='p4'></a>

## <span style='color:Red'>  QUESTION 4: Evidence for Money Ball

In the book [MoneyBall](https://en.wikipedia.org/wiki/Moneyball), Michael Lewis documented the introduction of statistics and data science in selecting players for the Oakland A's. in this problem, we're going to use two datasets from the [Sean Lahman's Baseball Database](http://seanlahman.com/) which contains the "complete batting and pitching statistics from 1871 to 2022, plus fielding statistics, standings, team stats, managerial records, post-season data, and more."   

For more details about this data, see documentation.txt in the `data` folder.  




## Exploratory Data Analysis

Throughout this question you will be asked to explore the dataset.  
You will need familiarity with all the methods and attributes introduced in  Chapter 6 of the assigned reading (and Lessons 2 and 3 of Lecture):  [Learning Data Science Chapter 6 - All (Subsetting, Aggregating, Joining, Transforming)](https://learningds.org/ch/06/pandas_intro.html).  

There is frequently more than one way to answer question in Pandas, so we encourage you to dive in and play around with different approaches.  Unless otherwise indicated, you are welcome to use whatever methods/attributes/approach that you choose.

There are two .csv files in the `data` folder of this HW assignment:  `teams.csv` and `salaries.csv`

Import the `Teams.csv` data into a DataFrame called `teams_df` and take a look at the data. 

Import the `Salaries.csv` data into a DataFrame called `salaries_df` and take a look at the data.


```python
teams_df= ...
teams_df
```

```python
salaries_df= ...
salaries_df
```

<!-- BEGIN QUESTION -->

### Question 4a). EDA:  Structure, Granularity and Faithfulness

Examine the structure, granularity and faithfulness of the datasets.   (Hint: The common utility functions we covered in lecture will be useful here).  

Then answer the following questions:

 - i).  What does the column `G` represent in the teams dataset?  (For a description of the columns, see the documentation in the `data` folder).  

 - ii).  What is the granularity of the `teams.csv` file?  

 - iii).  What is the granularity of the `salary.csv` file?
 
 - iv).  How many rows and columns are in the teams dataset?    Assign your answer to the variables `team_rows` and `team_col` below.  

 - v).  How many rows and columns are in the salary dataset? Assign your answer to the variables `salary_rows` and `salary_col` below.

 - vi).  How many entries in the `teams.csv` file are missing Attendance Data?  Assign your answer to the variable `missing_attendance` below.

### Answer Cell for Questions 4a(i)(ii)(iii) 
In this cell, answer questions 4a(i) - (iii) using Markdown (not code).  

**4a(i) Answer:**

**4a(ii) Answer**:

**4a(iii) Answer**:

In the code cells below justify your answers to part (ii) and (iii) and then answer parts iv through vi

```python
...

# Show work in this cell justifying your answer to part 4a(ii) (hint: either use .value_counts() or .groupby)
```

```python
...

# Show work in this cell justifying your answer to part 4a(iii) (hint: either use .value_counts() or .groupby)
```

```python
# Solution Cell for 4a(iv) and 4a(v)
#Use code to find the number of rows and columns in the teams data.  Do not enter any values by hand.

team_rows= ...

team_col = ...

salary_rows = ...

salary_col = ...


```

```python
# Solution Cell for 4a(vi)
# Use code to find the number of rows in the Teams data that are missing attendance data

missing_attendance = ...

```

```python
grader.check("q4a")
```

<!-- END QUESTION -->

### Question 4b) EDA:  Examining Scope and Temporality of the Data
 Use code to explore the datasets and answer the questions below:
 
 i).  How many unique teams are represented in the `teams_df` DataFrame?  Write code to calculate this and assign the answer to the variable `unique_teams`   
 
 ii). How many unique teams are represented in the `salaries_df` DataFrame? Write code to calculate this and assign the answer to `unique_teams_salary`
 
 iii).  What years does the `teams_df` data span?  
 Write code to calculate this and assign the answers to `teams_startyear` and `teams_endyear`

 
 iv).  What years does the `salaries_df` data span?   
 Write code to calculate this and assign the answers to `salaries_startyear` and `salaries_endyear`
 

```python
unique_teams = ...

#print out result
unique_teams

```

```python
unique_teams_salary= ...

#print out result
unique_teams_salary
```

```python
teams_startyear = ...

teams_endyear = ...

#print out result
[teams_startyear, teams_endyear]
```

```python
salaries_startyear = ...

salaries_endyear = ...

#print out result
[salaries_startyear, salaries_endyear]
```

```python
grader.check("q4b")
```

### Question 4c: 

If you haven't already, **read the assigned textbook readings** [on Filtering](https://learningds.org/ch/06/pandas_subsetting.html#filtering-rows)  [on Extracting Data](https://learningds.org/ch/06/pandas_subsetting.html) and [on Transforming Data](https://learningds.org/ch/06/pandas_transforming.html)

 - i ). For the data science approach used by the Oakland A's team in 2002, they used historical baseball data starting in 1998.    Filter the `teams_df` to create a new dataframe `teams_df_moneyball` with index `teamID` and columns `yearID`, `name`, `G`, `W`.  The new dataframe should only include records where the year is greater than or equal to 1998.
 
 - ii).  Add a new column called `winFrac` to the dataframe `teams_df_moneyball` that calculates the number of Wins divided by the number of Games for a given team in a given year, **rounded to the nearest hundredth**.  (Hint: Don't use any for loops).  


 - iii).  Use an appropriate extraction operator to determine the Oakland A's ("OAK") fraction of wins in the year 1998 (i.e. write code below  that selects this value from the `teams_df_moneyball` table) and assign it to the variable `winFrac_OAK_98`.   
  
 


```python
# Create teams_df_moneyball dataframe


teams_df_moneyball = ...


```

```python
# Add column called winFrac to the dataframe


...



# Print out resulting DataFrame
teams_df_moneyball.head()

```

```python
# Answer to part 4ci:

winFrac_OAK_98= ...

winFrac_OAK_98
```

```python
grader.check("q4c")
```

### EDA:  Visualizing Data 

Let's examine how winFrac has varied over time for the Oakland A's.

```python
#The code below uses the module plotlyexpress, which we imported at the top of the notebook using alias px
#Read the code below to see how to use this function, then run this cell

px.scatter(teams_df_moneyball.loc["OAK"],x="yearID", y="winFrac")
```

<!-- BEGIN QUESTION -->



Let's compare Oakland's winFrac to 2 other teams:  Boston Red Sox and the New York Yankees:

```python
#Read the entries below to see how to create this plot:
#Then run this cell

px.scatter(teams_df_moneyball.loc[["OAK","NYA", "BOS"]],x="yearID", y="winFrac",facet_col = 'name')
```

**Observations:**
In the cell below write down at least 3 observations you can make from these plots.


_Type your answer here, replacing this text._

<!-- END QUESTION -->



**Payrolls Per Team**  An interesting part of this story also takes into account the salary differences between players on these teams.   We'll examine this further in lecture and in our TA-led discussion!



<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 2!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 


* **Homework 2 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW2 Autograded assignment in Gradescope.

* **Homework 2 Manually Graded**: Submit your hw02.PDF to the HW2 Manually Graded assignment in Gradescope.  **It is your responsibility to fully review your PDF file before submitting and make sure that all your lines of code are visible and any LaTeX has correctly compiled and is fully viewable.**  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.** If not, you will lose points.    

**TROUBLESHOOTING TIPS**  If you are having any issues compiling your assignment, please [read through these troubleshooting tips first](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing), then post any questions on Piazza.  

**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw02.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for 3a (points: 3)
```
>>> assert get_hash(round(p_HHHTHT,6)) == '8621a84afb6956533591c0832ff1ba09'
```


Tests for 3c (points: 1)
```
>>> assert get_hash(p_hat)=='54fbf38cf649866815e0fefc46a1f6c7'
```


Tests for q1a (points: 3)
```
>>> assert get_hash(int(q1a_answer))=='da4fb5c6e93e74d3df8527599fa62642'
>>>
```


Tests for q1bcheck (points: 0)
```
>>> assert get_hash(round(q1b_answer,6)) == '0f7087f36d3e157843feb7566cb4d407'
```


Tests for q1ccheck (points: 0)
```
>>> assert get_hash(round(q1c_answer,6))== 'caed457b2e3a7d5fa3ede3de63e1fe34'
```


Tests for q4a (points: 5)
```
>>> # TEST
>>> assert get_hash(team_rows)== '51be2fed6c55f5aa0c16ff14c140b187'
```
```
>>> # TEST
>>> 
>>> assert get_hash(team_col)== '642e92efb79421734881b53e1e1b18b6'
```
```
>>> # TEST
>>> assert get_hash(salary_rows) == '398a516d8fbfdd9158eec76b283dd9ce'
```
```
>>> # TEST
>>> assert get_hash(salary_col) == 'e4da3b7fbbce2345d7772b0674a318d5'
```
```
>>> # TEST
>>> assert get_hash(missing_attendance) == 'd395771085aab05244a4fb8fd91bf4ee'
>>>
```


Tests for q4b (points: 4)
```
>>> # TEST
>>> assert get_hash(unique_teams) == 'f2217062e9a397a1dca429e7d70bc6ca'
```
```
>>> # TEST
>>> assert get_hash(unique_teams_salary) == '1c383cd30b7c298ab50293adfecb7b18'
```
```
>>> # TEST
>>> 
>>> assert get_hash(teams_startyear) ==  'fb8feff253bb6c834deb61ec76baa893'
```
```
>>> # TEST
>>> 
>>> assert get_hash(teams_endyear) == '3a824154b16ed7dab899bf000b80eeee'
```
```
>>> # TEST
>>> assert get_hash(salaries_startyear) == '1f36c15d6a3d18d52e8d493bc8187cb9'
```
```
>>> assert get_hash(salaries_endyear) == '95192c98732387165bf8e396c0f2dad2'
```


Tests for q4c (points: 4)
```
>>> # TEST
>>> assert get_hash(winFrac_OAK_98)== '3b97a72e9bc7f5c5e225aec9b0bf548c'
```
```
>>> # TEST
>>> assert get_hash(teams_df_moneyball.shape[1]) == 'e4da3b7fbbce2345d7772b0674a318d5'
```
```
>>> # TEST
>>> assert get_hash(teams_df_moneyball.shape[0]) == 'b137fdd1f79d56c7edf3365fea7520f2'
```
```
>>> # TEST 
>>> assert get_hash(teams_df_moneyball.query("yearID==2010").loc["BOS","winFrac"]) == '10efc2614821dcf1b1ebc941ebd6ce1f'
```
