```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

# Homework 7: The Central Limit Theorem & Intro to Testing Hypotheses


**Recommended Readings**: 

* [The Central Limit Theorem](https://inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html)

* [Testing Hypotheses](https://www.inferentialthinking.com/chapters/11/Testing_Hypotheses.html)




## Due Date: Thursday March 7th by 11:59 PM MT on Gradescope


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Piazza!


If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced

#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1 | 6 | manual
Question 2 | 6 | autograded
Question 3 | 16 | manual
Question 4abcdefgij | 18 | autograded
Question 4h | 4 | manual
|Total | 50

```python
import numpy as np
import pandas as pd
import math
import scipy.stats as stats
from scipy import special
import matplotlib
import matplotlib.pyplot as plt
from matplotlib.ticker import PercentFormatter
import seaborn as sns
import hashlib
from IPython.display import display, Latex, Markdown

plt.style.use('fivethirtyeight')
# Default plot configurations
plt.rcParams['figure.figsize'] = (16,8)
plt.rcParams['figure.dpi'] = 150
sns.set()
import warnings
warnings.simplefilter('ignore', FutureWarning)


def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

%matplotlib inline

```

<!-- BEGIN QUESTION -->

## Question 1

*** 

The probability density function for a *normal distribution with mean $\mu$ and SD $\sigma$* is defined by

$$f(x) ~ = ~ \frac{1}{\sigma\sqrt{2\pi}} e^{-\frac{1}{2}\big{(}\frac{x-\mu}{\sigma}\big{)}^2}, ~~~ -\infty < x < \infty$$

The following picture depicts a much-often spouted fact in statistics classes that roughly $68\%$ of the probability for a normal distribution falls within 1 standard deviation of the mean, roughly $95\%$ falls within two standard deviations of the mean, etc:

<img src="normal.jpeg" width="450px">


**Use Calculus to prove that the inflection point(s) of the probability density function of the normal distribution occur at $x = \mu \pm \sigma$** .  You must do the derivatives by hand (i.e. do not plug them into an online calculator) and show all steps.     Once you find all candidates for inflection points you can stop (you do not need to test values on either sides of the inflection points to show that the concavity changes, as we can see this visually in the graph of the pdf above).

Show all steps using LaTeX in the Markdown cell below:


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Question 2 :


A hardware store receives a shipment of 10,000 bolts that are supposed to be 12 cm long. 
The mean of this shipment of 10,000 bolts is indeed 12 cm, and the standard deviation is 0.2 cm.   

For the following questions, determine if you have enough information to answer.  If you do, then use built-in python functions to calculate the probability. If not, enter -1 for your answer.


**Question 2a)**. What is the probability that a randomly chosen bolt is less than 10 cm long?
 


```python
q2a_answer = ...
```

```python
grader.check("q2a")
```

**Question 2b)**.  For quality control, the hardware store chooses 100 bolts at random to measure.  They will declare the shipment defective and return it to the manufacturer if the average length of 100 bolts is less than 11.97 cm or greater than 12.04 cm.    Find the probability that the shipment is found satisfactory (i.e. not defective).

```python
q2b_answer = ...
```

```python
grader.check("q2b")
```

## Question 3
*** 

The file `income_data.csv` contains Age and Income information from a fictitious city in the Midwest with 5000 residents.   Consider this data to be our population data.



<!-- BEGIN QUESTION -->

**QUESTION 3A**: Load the data into a pandas DataFrame called `dfIncome`, calculate the population mean of Income and then  make a **density histogram** of the Distribution of the Income data **with 15 bins**.   Include a title for your plot and label the x-axis (we have provided a label for the y-axis).
Note we have included code to mark where the population mean lies on the histogram.

```python
dfIncome = ...

mean_income = ...

print("Population income mean is", mean_income)

...
# write code to plot density histogram above this line

plt.ylabel("Probability Per Dollar") # Since this is a density histogram, the y-units are probability per unit x
#Add a triangle marker to indicate where the population mean is
plt.scatter(mean_income, -.0000001, marker='^', color='red', s=300)

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**QUESTION 3B**: Describe the shape of the Income distribution (i.e. comment on modality and skew)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**QUESTION 3C**: 

i).  Write a function to collect a random sample of size `sample_size` with replacement from `dfIncome`  and plot the **density histogram** of the empirical distribution of the income for your sample.  Use 15 bins for your histogram and set the x-axis range to be from (0,210000).   (Hint: use the dataframe method `.sample()`). Include a title and label for both axes.          
Then run the cells provided below to output 3 separate distributions for sample sizes of $10$, $100$ and $1000$. 


```python
def income_sample(df, sample_size):
    
...
# rest of code above this line
```

```python
income_sample(dfIncome,10)
```

```python
income_sample(dfIncome,100)
```

```python
income_sample(dfIncome,1000)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Part 3cii).  What happens to the shape of the empirical sample distributions of income as you increase the sample size? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**QUESTION 3D:**  


If we want to estimate the **mean** of the population we can draw a sample from the population and compute the sample mean.  As we learned in class, since samples can vary, the sample mean can vary and thus it is a random variable and has its own distribution.


i).  Complete the function `income_sample_mean` below to randomly sample `sample_size` rows from `dfIncome` with replacement and return the sample mean of income for that sample. 


ii).  Complete the function `income_sample_dist` below to simulate `num_simulations` of randomly sampling `sample_size` rows from `dfIncome` with replacement and calculate the sample mean of income for each sample.
Store the sample means in an np.array called `means`.  The function should output a **density** histogram of the empirical sample mean income distribution.    On the histogram, include two markers on the histogram: A red one for the population mean (that you calculated in part 3A) and a yellow one for the mean of the `num_simulations` sample mean estimates.
Include a title and labels for the x and y-axis.

Then run the cells provided below to output 3 separate distributions for `num_simulations`=1000 and `sample_size` = $10$, $100$ and $1000$



```python
def income_sample_mean(df, sample_size):
    ...
```

```python
def income_sample_dist(df, sample_size, num_simulations):
    means = ...
      # 'means' stores "num_simulations" means from samples of size "sample_size"
    plt.xlim([20000,120000])
    ...
    
# Your code for part (ii) above
```

```python
income_sample_dist(dfIncome, 10, 1000)
```

```python
income_sample_dist(dfIncome, 100, 1000)
```

```python
income_sample_dist(dfIncome, 1000, 1000)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**QUESTION 3E:** 

Describe the shapes of the empirical sample mean distributions (comment on their modality and skew compared to the modality and skew of the population distribution).
What happens to the mean and standard deviations of these distributions as you increase the sample size?
What is the name of the theorem that explains what you are observing?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## QUESTION 4:  

The Therapeutic Touch (TT) is the idea that everyone can feel the Human Energy Field (HEF) around individuals.  Those who practice TT have described different people's HEFs as "warm as Jell-O" and "tactile as taffy." 

TT was a popular technique used throughout the 20th century that was toted as a great way to bring balance to a person's health. Certain practitioners claim they have the ability to feel the HEF and can massage it in order to promote health and relaxation in individuals.

### Emily Rosa

[Emily Rosa](https://en.wikipedia.org/wiki/Emily_Rosa) was a 4th grade student from Loveland, Colorado who was very familiar with the world of TT, thanks to her parents, who were both medical practitioners and skeptics of TT.

For her 4th grade science fair project, Emily decided to test whether or not TT practitioners could truly interact with a person's HEF. She later went on to publish her work in TT, becoming the youngest person to have a research paper published in a peer reviewed medical journal.

### Emily's Experiment

Emily's experiment was clean, simple, and effective. Due to her parents' occupations in the medical field, she had wide access to people who claimed to be TT practitioners. 

Emily took 21 TT practitioners and used them for her science experiment. She would take a TT practitioner and ask them to extend their hands through a screen (which they can't see through). Emily would be on the other side and would flip a fair coin. Depending on how the coin landed, she would put out either her left hand or her right hand. The TT practitioner would then have to answer which hand Emily put out. If a pracitioner could truly interact with a person's HEF, it would be expected that they answered correctly.

Overall, through 210 samples, the practitioner picked the correct hand 44% of the time. 

Emily's main goal here was to test whether or not the TT practitioners' guesses were random, like the flip of a coin. In most medical experiments, this is the norm. We want to test whether or not the treatment has **any** effect, *not* whether or not the treatment actually worked.
We will now begin to formulate this experiment in terms of the terminology we learned in this course. 

**QUESTION 4A:**  Which of the following is Emily's null hypothesis? 

 - (a):  Practitioners are answering independently and better than if they were guessing at random (i.e. the probability of any practitioner correctly answering is greater than 50%)
 
 - (b): Practitioners are answering independently and randomly (i.e. the probability of any practitioner correctly answering is 50%).  
  
 - (c):  Practitioners are answering independently and worse than if they were guessing at random (i.e. the probability of any practitioner correctly answering is less than 50%)
 
Assign `null_hypothesis` to the correct letter (enter it as a lower case string, for example 'a').  Note that the test below just checks that you entered a valid letter, it doesn't check if the letter is correct (that will be done by a test when you submit the notebook).




*Note:* Check out [11.3](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#step-1-the-hypotheses) for a refresher on hypotheses.


```python
null_hypothesis = ...
```

```python
grader.check("q4a")
```

**QUESTION 4B:** Given that we want to test whether or not the treatment has **any** effect, *not* whether or not the treatment actually worked, what is the alternative hypothesis?

Assign `alt_hypothesis` to the correct letter (enter it as a lower case string)

 - (a):  Practitioners are answering independently and better than if they were guessing at random (i.e. the probability of any practitioner correctly answering is greater than 50%)
 
  
 - (b):  Practitioners are answering independently and worse than if they were guessing at random (i.e. the probability of any practitioner correctly answering is less than 50%)
 
 - (c): Practitioners are answering independently and not randomly (i.e. the probability of any practitioner correctly answering is not 50%).  






```python
alt_hypothesis = ...
```

```python
grader.check("q4b")
```

**QUESTION 4C:** 
Remember that the practitioner got the correct answer 44% (0.44) of the time. According to Emily's **null hypothesis**, on average, what proportion of times do we expect the practitioner to guess the correct hand? Make sure your answer is a number between 0 and 1. (Again, note that the test below does not check for the correct answer, it just checks that your answer is in the correct format). 


```python
expected_proportion_correct = ...
expected_proportion_correct
```

```python
grader.check("q4c")
```

**QUESTION 4D:** 

The goal now is to see if our deviation from this expected proportion of correct answers is due to something other than chance. 


We usually use a statistic to help determine which model the evidence points towards. One of the test statistics below is appropriate for testing the null and alternative hypotheses you indicated above. Assign the variable `valid_test_stat` to the letter  (lower case  string) corresponding to the correct test statistic for this scenario. 



 - (a): 50 minus the actual percent correct
 - (b): the actual percent correct minus 50 
 - (c): 44 minus the actual percent correct
 - (d): the actual percent correct minus 44
 - (e): the sum of 50 and the actual percent correct
 - (f): the absolute value of the difference between 50 and the actual percent correct
 - (g): the absolute value of the difference between 44 and the actual percent correct


```python
valid_test_stat = ...
valid_test_stat
```

```python
grader.check("q4d")
```

**QUESTION 4E:**  Define the function `statistic` which takes in an expected proportion and an actual proportion, and returns the value of the statistic chosen in Question 4d. Assume that the argument takes in proportions, but  **return your answer as a percentage**   (i.e. an answer of 0.70 should be output as $70$)

*Hint:* Remember we are asking for a **percentage**, not a proportion. 


```python
def statistic(expected_prop, actual_prop):
    ...
```

```python
grader.check("q4e")
```

**QUESTION 4F:** Use your newly defined function from 4e to calculate the observed statistic from Emily's experiment. 

```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q4f")
```

**QUESTION 4G:**  Next we want to explore if this observed statistic is consistent with what we expect to see under Emily’s null hypothesis.  To do this, we will need to simulate test statistics under the null hypothesis.  

Let $X$ be the number of times the practioners are correct out of $210$ trials assuming the null hypothesis is true.  What distribution does $X$ have?

Assign the variable `X_dist` to the letter  (lower case  string) corresponding to the correct distribution for this scenario. 


 - (a): $X \sim Bin(210, 0.44)$
 - (b): $X \sim Bin(210, 0.5)$
 - (c): $X \sim Ber(.5)$
 - (d): $X \sim Ber(.44)$
 - (e): $X \sim Exp(.44)$
 - (f): $X \sim Exp(.55)$


```python
X_dist = ...
```

```python
grader.check("q4_g")
```

<!-- BEGIN QUESTION -->

**QUESTION 4H.** 

Create an array called `simulated_statistics` that contains 50,000 simulated values of the test statistic under the null hypothesis. Assume that the original sample consisted of 210 experiments. 

As usual, start by defining a function `one_simulated_statistic()` that simulates one value of the statistic. Your function should use `np.random.DISTRIBUTION` where `DISTRIBUTION` is the distribution you chose in part 4g.   Your function should also use your `statistic` function from part 4e.  


We have included the code that plots the distribution of the simulated values. The red dot represents the observed statistic you found in Question 4f.


```python
def one_simulated_statistic():
    ...

num_simulations = 50000

simulated_statistics = ...


# Run the this cell a few times to see how the simulated statistic changes
one_simulated_statistic()
```

```python
# Run this cell to produce a histogram of the simulated statistics
plt.hist(simulated_statistics, density = True, ec= "white")
plt.xlabel('Simulated Statistic')
plt.ylabel('Percent per Unit')
plt.title('Histogram of Simulated Statistics')
plt.gca().yaxis.set_major_formatter(PercentFormatter(1))
plt.scatter(observed_statistic, -0.002, color='red', s=100);
plt.show()
```

<!-- END QUESTION -->

**QUESTION 4I.** Using `simulated_statistics`, `observed_statistic`, and `num_simulations`, find the empirical p-value based on your simulation. 

```python
p_value = ...
p_value
```

```python
grader.check("q4_i")
```

**Question 4j** 


What is the conclusion of this hypothesis test using the 5% cutoff for the p-value?

 - (a). The data leads us to reject Emily's null hypothesis that the practitioners were just guessing.  This indicates that something more than just random chance is at play. 

 - (b). Given the data, we fail to reject Emily's null hypothesis that the practitioners are answering independently and randomly (i.e. we fail to reject that the probability of any practitioner correctly answering is 50%).  This doesn't prove that her hypothesis is correct, it just says we didn't find strong enough evidence to reject it.

 - (c) We accept Emily's null hypothesis that the probability of any practitioner correctly answering is 50%.  This proves that she was correct.




```python
conclusion = ...
```

```python
grader.check("q4_j")
```

Therapeutic touch fell out of use after this experiment, which was eventually accepted into one of the premier medical journals. TT practitioners hit back and accused Emily and her family of tampering with the results, while some claimed that Emily's bad spiritual mood towards therapeutic touch made it difficult to read her HEF. Whatever it may be, Emily's experiment is a classic example about how anyone, with the right resources, can test anything they want!


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 7!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 7 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW7 Autograded assignment in Gradescope.

* **Homework 7 Manually Graded**: Submit your hw07.PDF to the HW7 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw07.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2a (points: 3)
```
>>> assert get_hash(q2a_answer) == '6bb61e3b7bce0931da574d19d1d82c88'
```


Tests for q2b (points: 3)
```
>>> # TEST
>>> assert get_hash(round(q2b_answer, 3)) == '0a4a2fbfa050681b56c8bfce230ee97c'
```


Tests for q4_g (points: [0])
```
>>> # TEST
>>> assert X_dist in ['a', 'b', 'c', 'd','e', 'f']
```


Tests for q4_i (points: [0, 0])
```
>>> import numpy as np
>>> assert (type(p_value) == float or type(p_value) == np.float64)
```
```
>>> assert (0 <= p_value <= 1)
```


Tests for q4_j (points: [0, 2])
```
>>> # TEST
>>> assert conclusion in ['a', 'b', 'c']
```
```
>>> assert conclusion == 'b'
```


Tests for q4a (points: [0])
```
>>> # TEST
>>> assert null_hypothesis in ['a', 'b', 'c']
```


Tests for q4b (points: [0])
```
>>> # TEST
>>> assert alt_hypothesis in ['a', 'b', 'c']
```


Tests for q4c (points: [0])
```
>>> # TEST
>>> 
>>> assert 0 <= expected_proportion_correct <= 1
```


Tests for q4d (points: [0])
```
>>> # TEST
>>> assert valid_test_stat in ['a', 'b', 'c', 'd','e', 'f', 'g']
```


Tests for q4e (points: 2)
```
>>> assert int(round(statistic(.5,.5) + statistic(.4,.1),1)) == 30
```
```
>>> assert int(statistic(.4,.1) - statistic(.1,.4)) == 0
```


Tests for q4f (points: 1)
```
>>> assert int(round(observed_statistic,2))==6
```
