```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

# Homework 4

## Due Date: Thursday, Feb 8 by 11:59 PM MT on Gradescope


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Piazza!


If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced

#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1 | 5 | autograded
Question 2 | 4 | manual
Question 3 | 7 | manual
Question 4a | 4 | manual
Question 4bi-iii | 6 | autograded
Question 4biii | 4 | manual
Question 5 | 9 | manual
Question 6 | 11 | manual
|Total | 50

## Recommended Readings for this HW:

- [Probability of Equally Likely Outcomes](https://chrispiech.github.io/probabilityForComputerScientists/en/part1/equally_likely/)
- [Law of Total Probability](https://chrispiech.github.io/probabilityForComputerScientists/en/part1/law_total/)
- [Bayes Theorem](https://chrispiech.github.io/probabilityForComputerScientists/en/part1/bayes_theorem/)

- [Computational and Inferential Thinking: Section 9.5 - Finding Probabilities](https://inferentialthinking.com/chapters/09/5/Finding_Probabilities.html)
- [Computational and Inferential Thinking: Chapter 18: Updating Predictions](https://inferentialthinking.com/chapters/18/Updating_Predictions.html)
- [Computational and Inferential Thinking: Section 9.3 - Simulation](https://learningds.org/ch/10/eda_intro.html)
- [Theory Meets Data: Chapters 1, 2 and Sections 3.1 and 3.2](http://stat88.org/textbook/content/Chapter_01/00_The_Basics.html) 




<a id='top'></a>
---
**Shortcuts:**  [Problem 2](#p2) | [Problem 3](#p3) | [Problem 4](#p4) | [Problem 5](#p5) |  [Problem 6](#p6) 
---
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />



In this homework you'll investigate probability. You'll also use probability simulations to estimate probabilities, gain intuition about random processes and to check your pencil and paper work.   


```python
import pandas as pd
import numpy as np
import hashlib
import matplotlib.pylab as plt
%matplotlib inline


def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

### Problem 1: Roulette!


A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. Below is an example of a roulette wheel!

<img src="img/roulette_wheel.jpeg" width="330px">


On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green. 

**In each part below, write code that evaluates to the chance of the event described.**

**Part 1A (1 pt).**  What is the probability the winning pocket is black on all of the first three spins?


```python
first_three_black = ...
```

```python
grader.check("q1a")
```

**Part 1B (1 pt)**  What is the probability the color green never wins in the first 10 spins?

```python
no_green = ...
```

```python
grader.check("q1b")
```

**Part 1C (1 pt)** What is the probability the color green wins at least once on the first 10 spins?


```python
at_least_one_green = ...
```

```python
grader.check("q1c")
```

**Part 1D (2 pts).** What is the probability two of the three colors never win in the first 10 spins?



```python
lone_winners = ...
```

```python
grader.check("q14")
```

<!-- BEGIN QUESTION -->

***

[Back to top](#top)

<a id='p2'></a>

### Problem 2

Let $A$ and $B$ be events in a sample space $\Omega$.  

Suppose that the probability that $A$ occurs is $0.2$, the probability that $B$ occurs is $0.6$, and the probability that neither $A$ nor $B$ occur is $0.3$ (that is, the joint probability that $A$ doesn't occur AND $B$ doesn't occur is  $0.3$). 

**Determine if you have enough informatin to answer the following questions.  If so, write your solution showing all steps.  If not, explain what additional information you would need.**


**2a) (2 pts)**  What is $P(A, B)$?

**2b) (2 pts)**  What is $P(B\mid A')$?

**Hint:  Use the information provided to create a joint probability table for $A$ and $B$.**

Write up your full solution to both questions in the SAME box below using LaTeX (not code).  Show all steps fully justifying your answers.  


_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for Problem 2a, enter the answer (just the probability) you came up with in the cell below.  If you said you didn't have enough information, enter the number 0.  Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the cell above.  

```python
q2a_answer = ...
```

```python
grader.check("q2acheck")
```

To check your answer for Problem 2b, enter the answer you came up with (just the probability) in the cell below. If you said you didn't have enough information, enter the number 0. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the Markdown cell provided above.  

```python
q2b_answer = ...
```

```python
grader.check("q2bcheck")
```

<!-- BEGIN QUESTION -->

[Back to top](#top)

<a id='p3'></a>

### Problem 3

The accuracy of a diagnostic test is often described using the following terms:

 - Test Sensitivity: Ability to detect a positive case (i.e. probability that the test is positive given that the person actually has the virus).

 - Test Specificity:  Ability to determine a negative case (i.e. the probability that a person tests negative given that they don't have the virus)).


Suppose a  diagnostic test for a virus is reported to have $90\%$ sensitivity and $92\%$ specificity.

Suppose $2\%$ of the population has the virus in question.  

Answer the following questions all in ONE cell below using LaTeX.  Show all steps.

**3a) (6 pts)**.  If a person is chosen at random from the population and the diagnostic test indicates that they have the virus, what is the conditional probability that they do, in fact, have the virus?    Write up your full solution using LaTeX.

**3b) (1 pt)**.  Terminology: What is the **prior** and what is the **likelihood** in this scenario?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for Problem 3a, enter the answer you came up with (just the probability) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer above.

```python
q3a_answer = ...
```

```python
grader.check("q3check")
```

<!-- BEGIN QUESTION -->

***
[Back to top](#top)

<a id='p4'></a>

### Problem 4: Poker!

A common example for discrete counting and probability questions are poker hands.  Consider using a standard 52-card playing deck, with card ranks `[A,2,3,4,5,6,7,8,9,10,J,Q,K]` across the standard 4 suits: `[C,D,H,S]`.

**Part 4A (4 pts)** 

Suppose we draw 5 cards at random from the deck without replacement.  

In Poker, **"Three of a Kind"** is defined as a hand that contains **three cards of one rank and two cards of two other ranks.**  Notice that in this definition, a **Full House** (a hand that contains three cards of one rank and two cards of another rank) is **NOT classified as "three of a kind".** https://en.wikipedia.org/wiki/List_of_poker_hands#Three_of_a_kind


<img src="img/three.png" width="330px">

**What is the probability of drawing 5 cards (without replacement) that are "three of a kind?"**
   

Typeset your work using LaTeX below.  Show work justifying all steps.  You may leave your answer in terms of a ratio of products, but you should simplify away any combinatoric notation such as ${n \choose k}$ or $P(n,k)$.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for 4a, enter the answer you came up with (just the probability) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer in the cell above.

```python
q4a_answer = ...
```

```python
grader.check("q4acheck")
```





One of the big payoffs of simulation is that it can let us answer some probability questions that are otherwise quite difficult.  We can instead just simulate the process and get approximate results based on simulation.



There are several ways to draw random samples using Python:


- `np.random.choice(a)` (https://numpy.org/doc/stable/reference/random/generated/numpy.random.choice.html) draws a random sample from a population whose elements are in an array `a`. The output is an array consisting of the sampled elements.

- `df.sample(n)` (https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.sample.html) draws a random sample of `n` rows from the DataFrame `df`. The output is a DataFrame consisting of the sampled rows. 

**Part 4B**  Poker with computers!

To simulate the scenario you calculated theoretically in part 4a, you will use the steps below to write code to simulate draws of a hand of 5 cards from a standard 52-card deck.  You'll repeat the simulation 10000 times, and report how many "three of a kind" hands were observed.  If your simulation code is correct, your answer should be very close to your theoretical answer from part 4a. 

**4bi)(2 pts)**.  
Use the Python lists `suits` and `ranks` provided below to create a Python `list` called `cards` with 52 strings that represnt a full deck of cards. 

In the string describing the card give the suit first and the rank second.  For example to represent the 3 of hearts use the string 'H3'. To represent the queen of diamonds use the string 'DQ'.   

Use loops and/or list comprehensions to create the `card` list - **do not enter in the values manually.** 
Your deck should have 52 distinct cards.  


```python
suits=['D','H','C','S']

ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', '10', 'J', 'Q', 'K']

...
# 4bi).  Write code to create a list called `cards` representing a 52-card deck above this line.

cards
```

```python
grader.check("q4i")
```

**4bii)(4 pts)**.  Define a function named `three_kind` to take a Python `list` of 5 cards as input (where the cards are given as strings such as `H3`) and check if the cards are Three of a Kind.  Your function should return `True` if the input is 3 of a kind and `False` otherwise. 

HINT: The function `np.unique` (https://numpy.org/doc/stable/reference/generated/numpy.unique.html) can make checking hands easier. 


```python
...
# 4bii).Write a function (above this line) 'three_kind' to check a list of 5 cards to see if it is 3 of a kind 
```

```python
grader.check("q4ii")
```

<!-- BEGIN QUESTION -->

**4biii)(4 pts)**. 

Write code in the space below that completes the following steps:

Step 1:  Write a function to simulate 10,000 random draws from `cards` of 5 cards each, and check if each draw is Three of a Kind. The function should return the overall proportion of random hands (out of the 10,000) in which Three of a Kind was observed. If you have coded your simulation correctly, your answer to this part should be very close to your theoretical answer from Part 4A.  

Step 2:  Let's visualize how this simulation converges to the theoretical probability.   In class, we plotted a running estimate of the probability of an event as a function of the number of trials in our simulation.  Write code that completes 10,000 random draws of 5 cards each, but this time outputs a plot of a running estimate of the proportion of hands that are Three of a Kind as a function of the number of trials (from 1 to 10,000) in your simulation.  **Include a red horizontal line on your plot with the theoretical probability that you calculated in part 4A.**  Be sure to include a title on your plot and be sure to label both your axes on the plot.






```python
...
#4biii).Write your code for Step 1 above this line
```

```python
hand=np.random.choice(cards, replace=False, size=5)

hand
```

```python
...
# 4biii) Write your code for Step 2 above this line
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
[Back to top](#top)

<a id='p5'></a>

### Problem 5

To play a game, you have a bag containing 28 fair [four-sided dice](https://en.wikipedia.org/wiki/Four-sided_die), with faces `{1,2,3,4}`.  This bag also contains 9 fair six-sided dice (faces `{1,2,3,4,5,6}`) and 3 fair twenty-sided dice (faces `{1,2,3,4,...,19,20}`).  Call these 3 classes of die "Four", "Six" and "Twenty" (or $D_4$, $D_6$, and $D_{20}$, for short).  You grab one die at random from the box.  

Work the following problems by hand and write up your full solution using LaTeX unless otherwise stated (but don't be afraid to simulate to check your result!). 

**Part 5A (3 pts)**: You grab one die at random from the box and roll it one time.  What is the probability of the event $R_5$, that you roll a 5? Explain your reasoning mathematically (using LaTeX).   

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for 5a, enter the answer you came up with (just the probability) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer in part 5a above.

```python
q5a_answer = ...
```

```python
grader.check("q5acheck")
```

<!-- BEGIN QUESTION -->

**Part 5B (3 pts)**: Suppose you roll a 5.  Given this information, what is the probability that the die you chose from the box is a Six-sided die? Write up your full solution using LaTeX. Show all steps.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for 5b, enter the answer you came up with (just the probability) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer in part 5b above.

```python
q5b_answer = ...
```

```python
grader.check("q5bcheck")
```

<!-- BEGIN QUESTION -->

**Part 5C (3 pts):**  Are the events $R_5$ and $D_6$ independent?  Write up your full solution using LaTeX. Show all steps. Justify your answer **using the mathematical definition of independence**.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for 5c, enter the answer you came up with as a string (e.g., `'in/dependent'`) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer in part 5c above.

```python
q5c_answer = ...
```

```python
grader.check("q5ccheck")
```

<!-- BEGIN QUESTION -->

***
[Back to top](#top)

<a id='p6'></a>

### Problem 6

Suppose you roll two fair six-sided dice.  Let $C$ be the event that the two rolls are _close_ to one another in value, in the sense that they're either equal or differ by only 1.  

**Part 6A (4 pts)**: Compute $P(C)$ by hand.  Show all steps using LaTeX.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for 6a, enter the answer you came up with (just the probability) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer in part 6a above.

```python
q6a_answer = ...
```

```python
grader.check("q6acheck")
```

<!-- BEGIN QUESTION -->

**Part 6B (3 pts)**: 
Write a simulation to run 10,000 trials of rolling a pair of dice and estimate the value of $P(C)$ you calculated in **Part A**.  Your estimate should agree with the exact calculation you did in **Part A**.  If it doesn't, try increasing the number of trials in your simulation. 

```python
...
#Your code above this line
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

##### **Part 6C (3 pts)**:  
In class we plotted a running estimate of the probability of an event as a function of the number of trials in our simulation.  Write code to run 5 independent simulations of 50,000 trials each to estimate $P(C)$ and plot their running estimate curves on the same set of axes. **Hint**: This is a lot of computation, so try to leverage Numpy and **list comprehensions** as much as possible so that your code doesn't run forever. 

**Include a red horizontal line on your plot with the theoretical probability that you calculated in part 6A.**  Be sure to include a title on your plot and be sure to label both your axes on the plot.

```python
...
# Your code above this line
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Part 6D (1 pt)**: Describe the behavior of the running estimates as the number of trials increases.

i).  What value(s) are they converging to?  
ii).  How many trials does it take until they appear to converge?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 4!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 4 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW4 Autograded assignment in Gradescope.

* **Homework 4 Manually Graded**: Submit your hw04.PDF to the HW4 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw04.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q14 (points: 2)
```
>>> assert get_hash(lone_winners) == '89c273b522bd98f03a71fb168b4bd51b'
```


Tests for q1a (points: 1)
```
>>> assert get_hash(first_three_black) == '51759fb527f916a54c86da4268f79f43'
```


Tests for q1b (points: 1)
```
>>> assert get_hash(no_green) == '537adeee371f1fa62c79f0788ebe1098'
```


Tests for q1c (points: 1)
```
>>> assert get_hash(at_least_one_green) == 'd0f2552c478594fecdb33230685f6e01'
```


Tests for q2acheck (points: 0)
```
>>> assert get_hash(q2a_answer) == 'cb5ae17636e975f9bf71ddf5bc542075'
```


Tests for q2bcheck (points: 0)
```
>>> assert get_hash(q2b_answer) == '3bd2a01cf8d7f2c39060b76f286628c8'
```


Tests for q3check (points: 0)
```
>>> assert get_hash(q3a_answer) == '782390db9c4a228e21ed5dc8a61d652e'
```


Tests for q4acheck (points: 0)
```
>>> assert get_hash(q4a_answer) == '0b99b9ebc73711aeb08047a09e3e65bb'
```


Tests for q4i (points: 2)
```
>>> assert get_hash(len(cards)) == '9a1158154dfa42caddbd0694a4e9bdc8'
```
```
>>> assert get_hash(sorted(set(cards))) == '83dfc5b12f2790365c2b2ba676e90de0'
```


Tests for q4ii (points: 4)
```
>>> assert (three_kind(['H3', 'S3', 'D3', 'H4', 'C5']) == True)
```
```
>>> assert (three_kind(['C4', 'D4', 'S4', 'H4', 'H2']) == False)
```
```
>>> assert (three_kind(['SA', 'S4', 'S4', 'H5', 'H10']) == False)
```
```
>>> assert (three_kind(['SA', 'DA', 'HA', 'C5', 'H5']) == False)
```
```
>>> assert (three_kind(['S3', 'D3', 'H4', 'C4', 'H5']) == False)
```
```
>>> assert (three_kind(['S10', 'D10', 'H10', 'CQ', 'HA']) == True)
```


Tests for q5acheck (points: 0)
```
>>> assert get_hash(q5a_answer) == '4e1856a66b8646915fecd9a23cb85c23'
```


Tests for q5bcheck (points: 0)
```
>>> assert get_hash(q5b_answer) == '90c55616ca660fd74557bad8cead21c3'
```


Tests for q5ccheck (points: 0)
```
>>> assert get_hash(q5c_answer.lower()) == '0aeb87dfc59607ef95459d79b98264f0'
```


Tests for q6acheck (points: 0)
```
>>> assert get_hash(q6a_answer) == '53f811b7284849904a1028d47c2ec903'
```
