```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Homework 5

## Due Date: Thursday, February 15th by 11:59 PM MT on Gradescope


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Piazza!


If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced

#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1 | 7 | autograded
Question 2acef | 11 | manual
Question 2bd | 5 | autograded
Question 3 | 10 | manual
Question 4abc | 3| autograded
Question 4defg | 14 | manual
|Total | 50


## Recommended Readings for this HW:

- [CS109 Course Reader: Random Variables](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/rvs/)
- [CS109 Course Reader: Probability Mass Functions (PMF)](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/pmf/)
- [CS109 Course Reader: Expectation](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/expectation/)
- [CS109 Course Reader: Variance](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/variance/)
- [CS109 Course Reader: Bernoulli Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/bernoulli/)
- [CS109 Course Reader: Binomial Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/binomial/)
- [CS109 Course Reader: Poisson Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/poisson/)
- [100 Binomial Question](https://chrispiech.github.io/probabilityForComputerScientists/en/examples/100_binomial_problems/)

<a id='top'></a>
---
**Shortcuts:**  [Problem 2](#p2) | [Problem 3](#p3) | [Problem 4](#p4) | 
---
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />



In this HW we explore properties of probability distributions, both theoretically and through simulation.

```python

import pandas as pd
import numpy as np
import hashlib
import matplotlib.pylab as plt
import time
%matplotlib inline


def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()


```

### Question 1

Factorials and the *binomial coefficients* $\binom{n}{k} = \frac{n!}{k!(n-k)!}$ get large very quickly as $n$ gets large. One way to compute them is to use the `SciPy` module `special`. `SciPy` is a collection of Python-based software for math, probability, statistics, science, and engineering.

```python
from scipy import special
```

Below are some examples of `special.factorial`:

```python
special.factorial(5)

special.factorial(range(1, 6))
```

Traditionally, subsets of $k$ individuals out of a population of $n$ individuals are called *combinations*, and so `special.comb(n, k)` evaluates to $\binom{n}{k}$. 

Note that we will always use the term *subsets* to mean un-ordered sets. We will use *permutations* in situations where we need to keep track of the order in which the elements appear.


Look at the code and output below carefully (including types) before starting Question 1.


https://docs.scipy.org/doc/scipy/reference/generated/scipy.special.comb.html



```python
special.comb(5, 3), special.factorial(5) / (special.factorial(3) * special.factorial(2))
```

```python
special.comb(5, range(6))
```

Consider a population in which a proportion $p$ of individuals are called "successes" (or 1, if you prefer) and the remaining proportion are rudely called "failures" (or 0).

As we saw in lecture, if you draw a sample of size $n$ (where $n$ is some positive integer) at random with replacement from the population, then the number of successes is a random variable that follows the binomial distribution. The probability of drawing $k$ successes and $n-k$ failures is $\binom{n}{k}p^k(1-p)^{n-k}$, for $0 \leq k \leq n$. To reduce writing, we will shorten "$k$ successes and $n-k$ failures" to "$k$ successes". In other words, "$k$ successes" means "exactly $k$ successes".

To formalize notation: if $X$ is the number of successes we draw, then $P(X = k) = \binom{n}{k}p^k(1-p)^{n-k}$.

### Question 1a) ###
Suppose you sample 100 times at random with replacement from a population in which 26% of the individuals are successes. Write a Python expression that evaluates to the chance that the sample has 20 successes. 

**Computational note:** Don't import any other libraries; just use the ones already imported and plug into the formula above. It's far from the best way numerically, but it is fine for the numbers involved in this Homework. 

```python
prob_1a = ...
prob_1a
```

```python
grader.check("q1a")
```

### Question 1b) ###
Complete the cell with a Python expression that evaluates to an array whose elements are the chances of $k$ successes for $k = 0, 1, 2, \ldots, 100$. That is, `all_probs[k]` should contain $P(X = k)$, where $X$ is the random variable describing the number of successes drawn from the scenario outlined above.

```python
k = np.arange(101)
all_probs = ...
np.round(all_probs, 4), sum(all_probs)
```

```python
grader.check("q1b")
```

Using the array of probability values we can plot a histogram of this discrete probability distribution as follows.

**NOTE:** We set the bin widths to be equal to 1 and add white lines between each value so that we can interpret area in the histogram as probabilities.

```python
k = np.arange(101)
p = special.comb(100, k)*(0.26**k)*(0.74**(100-k))


fig, ax = plt.subplots()

ax.bar(k, p, width=1, ec='white');
ax.set_axisbelow(True)
ax.grid(alpha=0.25)
plt.xlim(1,60)
#Always label your x-axis
plt.xlabel("k")
#Always label your y-axis
plt.ylabel("P(X=k)")

#Always include a title
plt.title("Binomial Distribution, n=100, p=0.26");


```

### Question 1c) ###


Complete the cell with an expression that evaluates to the chance that the number of successes in the sample is in the interval $26 \pm 10$ (inclusive on both sides). In other words, determine $P(16 \leq X \leq 36)$. **Hint: You might want to apply your Python indexing knowledge here.**

**Note: Please assign your answer to `prob_1c`. Because we use an autograder, please make sure your answer is a proportion between 0 and 1, not a percent between 0 and 100.**

```python
prob_1c = ...

```

```python
grader.check("q1c")
```

Consider (once again) a sample of size $n$ drawn at random with replacement from a population in which a proportion $p$ of the individuals are called successes.

Let $S$ be the random variable that denotes the number of successes in our sample. (As stated above, $S$ follows the binomial distribution.) Then, the probability that the number of successes in our sample is **at most** $s$ (where $0 \leq s \leq n$) is

$$P(S \leq s) = P(S = 0) + P(S = 1) + ... + P(S = s) = \sum_{k=0}^s \binom{n}{k}p^k(1-p)^{n-k}$$

We obtain this by summing the probability that the number of successes is exactly $k$, for each value of $k = 0, 1, 2, ..., s$.

### Question 1d) ###

Please fill in the function `prob_at_most` which takes $n$, $p$, and $s$ and returns $P(S \le s)$ as defined above. If the inputs are invalid: for instance, if $p > 1$ OR $s > n$ then return 0." 

It might help to refer to the calculation in Question 1b.

```python
def prob_at_most(n, p, s):
    """ 
    returns the probability of S <= s
    Input n: sample size; p : proportion; s: number of successes at most"""
    ...
    # Your code above this line
```

```python
grader.check("q1d")
```

### Question 1e) ###
In an election, supporters of Candidate C are in a minority. Only 45% of the voters in the population favor the candidate.

Suppose a survey organization takes a sample of 200 voters at random with replacement from this population. Use `prob_at_most` to write an expression that evaluates to the chance that a majority (more than half) of the sampled voters favor Candidate C.

```python
p_majority = ...
p_majority
```

```python
grader.check("q1e")
```

### Question 1f) ###
Suppose each of five survey organizations takes a sample of voters at random with replacement from the population of voters in Part **e**, independently of the samples drawn by the other organizations. 

- Three of the organizations use a sample size of 200
- One organization uses a sample size of 300
- One organization uses a sample size of 400

Write an expression that evaluates to the chance that in at least one of the five samples the majority of voters favor Candidate C. You can use any quantity or function defined earlier in this exercise.

```python
prob_1f = ...
prob_1f
```

```python
grader.check("q1f")
```

***

[Back to top](#top)

<a id='p2'></a>

### Question 2 - Simulating Roulette


A Las Vegas roulette wheel contains 38 numbers $\{0, 00, 1, 2, \ldots, 36\}$. Of the non-zero numbers, 18 are red and 18 are black. You can place bets on various number/color combinations and each type of bet pays-out at a different rate.  For example: 

- If you bet $\$1$ on red (or black) and win, then you win $\$1$ (i.e. you get your original dollar back, plus another dollar). 
- If you bet $\$1$ any particular number and win, then you win $\$35$ (i.e you get your original dollar back, plus $\$35$). 
- If you bet $\$1$ on the first dozen (or second dozen, or third dozen) nonzero numbers and win, then you win $\$2$ (i.e. you get your original dollar back, plus another $\$2$. 

<img src="img/roulette_wheel.jpeg" width="330px">

It seems like there are so many ways to win!  In reality, some very careful probability theory was done by the game designers to ensure that there is not much difference in any particular payout.  We'll explore roulette both by simulation and by hand in this exercise. 

<!-- BEGIN QUESTION -->

### Question 2a) ### 

Let $X$ be a random variable that gives your net winnings if you bet on red and the roulette wheel is spun once.   

i).  What is the probability distribution of $X$?  Give your answer as a table. 

ii).  Calculate the expected value of your net winnings by betting on red.   

iii). Calculate the variance of your net winnings by betting on red.

Write up your full solution to all 3 parts in the SAME box below using LaTeX (not code).  Show all steps fully justifying your answer.  


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2b) ### 

Let's simulate this.  In the first code box below, write code to simulate one spin of a roulette wheel. Your code should take in a single variable `num_spins` and return an array of roulette spins of that length **WITHOUT USING A LOOP**.  Your output should be an array of strings in the form of the number then the color (i.e. `array(['18R', '00G',...])`) 

In the 2nd code box below, write code that takes the number of spins and either the color red or black as input, calculates winnings for each spin assuming you bet on that color for all spins, and then outputs the average winnings out of those spins.  

Then run the simulation for `num_spins` = 10,000,000  and compare to your answer from part A.

**To receive credit you must write your code such that all lines are visible in your PDF output.**

```python
def spin_roulette(num_spins): 
    # Create wheel:
    red = [str(ii)+"R" for ii in [1,3,5,7,9,12,14,16,18,19,21,23,25,27,30,32,34,36]] 
    black = [str(ii)+"B" for ii in [2,4,6,8,10,11,13,15,17,20,22,24,26,28,29,31,33,35]]
    green = [str(ii)+"G" for ii in ["0", "00"]]

    wheel = red+black+green

    return np.random.choice(wheel, size=num_spins)

# Your code above this line

spin_roulette(5)

```

```python
def color_winnings(color='R', num_spins=100000):

...
# Your code above this line
    
print("E[Winnings] = {:.3f}".format(color_winnings(color="R", num_spins=int(1e7))))
```

```python
grader.check("q2b")
```

<!-- BEGIN QUESTION -->

### Question 2c) ### 

In Roulette you can bet on one of three "dozens" segments, called 1st 12, 2nd 12, and 3rd 12. They cover 1-12, 13-24, and 25-36, respectively.   If you bet $\$1$ on the first dozen (or second dozen, or third dozen) nonzero numbers and win, then you win $\$2$ (i.e. you get your original dollar back, plus another $\$2$. 


Let $Y$ be a random variable that gives your net winnings if you bet on any one of the three "dozen" nonzero numbers and the roulette wheel is spun once.  

i).  What is the probability distribution of $Y$?  Give your answer as a table. 

ii).  What is $E[Y]$?

iii). What is $Var[Y]$?

Write up your full solution to all 3 parts above in the SAME box below using LaTeX (not code).  Show all steps fully justifying your answer.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2d) ### 

Write code to simulate `num_spins` spins, record the winnings for each spin if you bet on the first dozen nonzero numbers, and calculate the average winnings out of the total spins.

Then run the simulation for `num_spins` = 10,000,000  and compare to your answer from part C.  


```python
def dozen_winnings(num_spins):

...
# Your code above this line
    
print("E[Winnings] = {:.3f}".format(dozen_winnings(num_spins=int(1e7))))

```

```python
grader.check("q2d")
```

<!-- BEGIN QUESTION -->

 ### Question 2e) ###  
Recall, we showed in class that the expected winnings if you bet on any number is also $-\dfrac{1}{19}$.    

So  you're hopefully onto the pattern by now.  The payouts in Roulette are designed so that the expected payout for a winning bet is always $-\dfrac{1}{19}$.  

Since we define these payouts in terms of your winnings after betting $1, we can think of these as payout odds.  

For example, since if you bet $\$1$ on the first dozen nonzero numbers and win, then you win $\$2$, we say the odds are 2 to 1 (denoted 2:1).   

The odds are 35:1 for landing on any particular number.  This means if you bet $\$1$, you’ll win $35.
 
 
Suppose the casino wanted to develop odds for a new bet in Roulette, where they allow you to bet on any set of 3 different numbers.   Let the odds for this new bet be $$x:1$$ What should $x$ be so that the expected payout for a winning bet is still $-\dfrac{1}{19}$?.

Show your work using LaTeX below.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

 ### Question 2f) ###  
Let's generalize this!   

Define a function $x(n)$ that describes the odds the casino should give for betting $\$1$ on any set of $n$ numbers if the casino wants to keep the expected payout for a winning bet at $-\frac{1}{19}$ for any $n$.   (For example, the odds for betting on any 3 different numbers should be set at $x(3)$ to $1$.  The odds for betting on any 4 different numbers should be set at $x(4)$ to $1$).   

Write your solution in the cell below:


_Type your answer here, replacing this text._

<!-- END QUESTION -->

***

[Back to top](#top)

<a id='p3'></a>

### Question 3


Consider the discrete random variable $X$, with probability distribution given by $P(X=k) = 2ak^2-2ak$ for $k = 2, 3, 4$, where $a$ is a constant.


<!-- BEGIN QUESTION -->

Answer all of the parts below in the SAME cell below using LaTeX.  Show all of your steps.  

**3a**).  Determine the value of $a$ such that this defines a valid probability distribution.  Use that value for the rest of the problem.  

**3b**).  Calculate $P(X\leq3)$.

**3c**).  What is $E[X]$? (Show steps calculating this).

**3d**).  What is the standard deviation of $X$? (Show all steps calculating this).

Answer all of the parts above in SINGLE cell provided below using LaTeX.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answers to 3c and 3d, enter them below.  Note that this autograder test is not worth any points, it's just for you to check your work above to make sure you're on the right track. 

```python
expected_value = ...
standard_deviation= ...
```

```python
grader.check("q3_check")
```

<!-- BEGIN QUESTION -->

### Question 3e

Plot a histogram of the discrete probability distribution for $X$.  

Use the same plotting guidelines as shown in Problem 1 so we can interpret area in the histogram as representing probability: 
 - Set the bin widths to be equal to 1 
 - Add white lines between each bar 

**Be sure to include a title and label your x and y axes on your plot.**


```python
...
# Your code for the histogram above this line

```

<!-- END QUESTION -->

***

[Back to top](#top)

<a id='p4'></a>


### Question 4


You are waiting in line at the grocery store. It is taking _forever_!  There are only two lines open; one is being tended by a cashier named [John Henry](https://en.wikipedia.org/wiki/John_Henry_(folklore)), and the other is tended by a [self check-out machine](https://theconversation.com/the-economics-of-self-service-checkouts-78593). Like all human beings when they arrive at the front of the store to check-out and encounter lines everywhere, you first experience a moment of intense panic. _Which line will be the fastest?_ you wonder, as people shuffle around you.

You decide you need to model the arrival of customers at the front of each of the lines.  If you haven't already, read the [required textbook reading on the Poisson Distribution].(https://chrispiech.github.io/probabilityForComputerScientists/en/part2/poisson/)  

From your Intro to Data Science class you remember that the distribution of independent random arrivals is often modeled using a Poisson distribution.  You observe the following:
* John's line checks-out an average of 4 customers per ten minutes,
* the self check-out machine checks-out an average of 5 customers per ten minutes **if** the machine is working properly, 
* the self check-out machine checks-out an average of 1 customer per ten minutes if the machine is freezing up, and
* in any given moment, the self check-out machine has a probability of 0.1 of freezing up.

Answer the following questions about this scenario. 

### Question 4a).  
What is the probability that exactly 6 customers pass through John's line in the next 10 minutes? 

```python
q4a = ...

q4a
```

```python
grader.check("q4a")
```

### Question 4b).  
What is the probability that exactly 6 customers pass through the self check-out in the next 10 minutes, given that it is working?  

```python
q4b = ...

q4b
```

```python
grader.check("q4b")
```

### Question 4c).
What is the probability that exactly 6 customers pass through the self check-out in the next 10 minutes, given that it is frozen? 


```python
q4c = ...

q4c
```

```python
grader.check("q4c")
```

<!-- BEGIN QUESTION -->

### Question 4d).  What is the probability that the self check-out tends exactly 6 customers in the next 10 minutes?
Show all steps in the cell below using LaTeX.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answers to 4, enter it below.  Note that to receive credit for your answers you must show work above,  that this autograder test is not worth any points, it's just for you to check your work above to make sure you're on the right track. 

```python
q4d_check = ...

q4d_check
```

```python
grader.check("q4d_check")
```

<!-- BEGIN QUESTION -->

### Question 4e) ###
S'pose John is working a 5-hour shift from 4-9 PM after school. He gets no breaks, because the year is 1870 and worker's rights is not yet a thing.  

Plot a histogram of the probability mass function (pmf) of the number of customers he serves in his 5 hour shift.   
For the **domain of the histogram, include $x$ values between $80$ and $160$ in your plot.**

Be sure to include labels for your x and y axes and a title for your plot.

**Hint:** Python has a built-in function to calculate the Poisson distribution for different values of $\lambda$.  See the documentation for `poisson.pmf` in `scipy.stats`  (https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.poisson.html)

**Hint:**  Since we are changing the time interval over which we are counting customers, you will need to update the parameter $\lambda$ in the Poisson distribution to be the average number of customers John can serve in a $5$-hour shift. **You can assume that his rate of 4 customers per 10 minutes scales up consistently during his 5 hour shift.**  



```python
from scipy.stats import poisson
```

```python
...

# your code above this line
```

<!-- END QUESTION -->

### Question 4f) ###

Calculate  the probability that John serves 100 or more customers during his **5-hour** shift.  Note for this problem do the calculation of the actual probability, not a simulation.   Assign your output to the variable `prob_100_or_more`.
Assume we can safely scale our data assumptions up to a 5-hour interval.



```python
...
# your code above this line
prob_100_or_more
```

```python
grader.check("q4f")
```

<!-- BEGIN QUESTION -->

### Question 4g) 

Time to simulate!

In Questions 4e and 4f above it was possible to calculate the probabilities exactly.  However, in some situations it may be too challenging to calculate an exact theoretical probability in which case we can use numerical simulation to estimate the probability.

In this part we're going to re-do Questions 4e and 4f via simulation to demonstrate that we're able to get very close to the theoretical probabilities via simulation.  

That is, we're going to simulate the distribution of the number of customers John serves in his 5 hour shift.

i).  Use the `np.random.poisson` function to randomly sample 100,000 times (with replacement) from a Poisson distribution with the same parameter $\lambda$ you used in Question 4e.   Save these simulated values in a `numpy array` called `random_sample`


ii). Then plot a **density histogram** of your simulated values.  Your histogram should look very similar to the theoretical histogram you plotted in Question 4e.  
Set the bin parameter to be `bins = np.arange(80,160)`.
Be sure to include a title and labels for the units of your x and y axes.


iii).  Then use your simulated values to calculate the (simulated) probability that John serves 100 or more customers in a 5-hour block. Your answer should be very close to the theoretical answer you calculated in part 4f.


```python
random_sample = ...

#Peek at the first 10 values
random_sample[:10]
```

```python
...
# Your code for part 4g(ii) above this line
```

```python
...

# Your code for part iii above this line

print('Simulated Probability that John serves 100 or more customers in a 5-hour block = {:.3f}'.format(p100_john))


# Output should match your theoretical answer to Part 4f
```

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 5!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 5 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW5 Autograded assignment in Gradescope.

* **Homework 5 Manually Graded**: Submit your hw05.PDF to the HW5 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw05.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw05.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 1)
```
>>> # TEST
>>> assert prob_1a >= 0 and prob_1a <= 1
```
```
>>> # TEST
>>> assert get_hash(round(prob_1a, 3)) == '5deac34732e7433a399ffd47b12b43f1'
```


Tests for q1b (points: 1)
```
>>> # TEST
>>> assert np.all(all_probs >= 0) and np.all(all_probs <= 1)
```
```
>>> # TEST
>>> k = np.arange(101)
>>> test_1b = special.comb(100, k)*(0.26**k)*(0.74**(100-k))
>>> assert np.all(np.isclose(test_1b, all_probs, atol = 1e-5))
```


Tests for q1c (points: 1)
```
>>> # TEST
>>> assert prob_1c >= 0 and prob_1c <= 1
```
```
>>> # TEST
>>> assert get_hash(round(prob_1c, 2)) == '5a01f7b0f3c45c7ecda1bdeb54968ce0'
```


Tests for q1d (points: 1)
```
>>> # TEST
>>> assert prob_at_most(3, 0.4, 1) >= 0
```
```
>>> # TEST
>>> assert prob_at_most(5, 0.6, 3) <= 1
```
```
>>> # TEST
>>> assert prob_at_most(2, 3, 4) == 0
```
```
>>> # TEST
>>> assert get_hash(round(prob_at_most(10, 0.7, 4), 3)) == '30e75abc2b9e6415a3ce76271d98835b'
```
```
>>> # TEST
>>> assert get_hash(round(prob_at_most(5, 1, 1), 1)) == '30565a8911a6bb487e3745c0ea3c8224'
```
```
>>> # TEST
>>> assert get_hash(np.round(prob_at_most(8, 0.7, 3), 5)) == 'e237b8ec306f62174dd05cfe5e9228a7'
```


Tests for q1e (points: 1)
```
>>> # TEST
>>> p_majority >= 0 and p_majority <= 1
True
```
```
>>> # TEST
>>> assert get_hash(round(p_majority, 3)) == 'd68a5a88e3ca0ee94ee5951c43fbb552'
```


Tests for q1f (points: 2)
```
>>> # TEST
>>> assert prob_1f >= 0 and prob_1f <= 1
```
```
>>> # TEST
>>> assert get_hash(round(prob_1f, 4)) == '22559430036a037ee145a651d1e91010'
```


Tests for q2b (points: 2)
```
>>> for spin in spin_roulette(10000):
...     assert spin[-1] in ("R","B","G")
...     assert (int(spin[:-1]) >= 0) and (int(spin[:-1]) <= 36)
```
```
>>> assert get_hash(round(color_winnings(color='R', num_spins=int(1e7)), 2)) == 'e927fe5f5ad45bcd256f7039b3783ffb'
```
```
>>> def timer(func, arg):
...     start = time.time()
...     func(arg)
...     return time.time() - start
>>> 
>>> assert timer(spin_roulette, int(1e7)) < 3
```


Tests for q2d (points: 3)
```
>>> assert get_hash(round(dozen_winnings(int(1e7)), 2)) == 'e927fe5f5ad45bcd256f7039b3783ffb'
```


Tests for q3_check (points: 0)
```
>>> # TEST
>>> assert get_hash(expected_value) == '804066a6a5042090a2a0b25c55f38e72'
```
```
>>> # TEST
>>> assert get_hash(round(standard_deviation,3)) == '9778132274276273b08f5bad04ebea24'
```


Tests for q4a (points: 1)
```
>>> # TEST
>>> 
>>> assert get_hash(round(q4a,4)) == '1b257ab8eb93c45e0c77f254700d71a1'
```


Tests for q4b (points: 1)
```
>>> # TEST
>>> 
>>> assert get_hash(round(q4b,4)) == 'f72b8e9036a7781c17983e1646b483b4'
```


Tests for q4c (points: 1)
```
>>> # TEST
>>> assert get_hash(round(q4c,5)) == 'ada08e891efaad3d5e003bab1b22d909'
```


Tests for q4d_check (points: 0)
```
>>> assert  get_hash(round(q4d_check,5)) == '3c0a26542c11e4e563ea650f65189940'
>>>
```


Tests for q4f (points: 3)
```
>>> # TEST
>>> 
>>> assert get_hash(round(prob_100_or_more,10))=='dd556b4f6319d0ea5dbe2002edcb53e8'
>>>
```
